Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and email below:

```python
# Full name
NAME = ""
# Institutional email (hm.edu or hmtm.de)
EMAIL = ""
```

---

# Rudiments of Python

**Why Python ?**

There are many reasons why we chose to teach you rudiments of Python in this tech crash course:

- It is free and Open Source,
- It is a worldwide language used in science and machine learning. For this reason, you will easily find support online, even for your specific problem, and it will continually be improved and updated,
- It has a simple syntax, which makes it easy to read and learn

## 1. Basic types

The main basic data types in python are:

- **Numbers** : those can be integers `int` or floating  `float`;
- **Strings** : a sequence of characters, noted `str`;
- **Booleans** : Either `True` or `False`, noted `bool`;

One function that is very useful to know is `type()`, which allows you to know the type of a variable.

```python
type(42)
```

```python
type(1.4)
```

```python
type("Hello World!")
```

```python
type(True)
```

Another important function is `print()`, which allows you to display a value. 

In Jupyter Notebook, you can display a value by simply using a cell, as shown above with `type("Hello world!)`. However,  if several variables are called in the same cell, only the last one will be displayed, except if you `print` them all !

As a rule of thumb, you should use `print()` to display values.

```python
42
"Hello world!"
```

```python
print(42)
print("Hello world!")
```

Depending of their types, values can be manipulated differently. 
Let's take the example of the `+` operator.

```python
type(2+2)
```

```python
type(2+2.0)
```

```python
2+'2'
```

<div class="alert alert-info">
<b>Instruction:</b> In the next cell, use the + operator with two strings.
</div>

<!-- grade_id: cell-9feded7ef9f10ecc, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1b9c781fd70f645f, points: 1, grade: true, solution: true -->
<div class="alert alert-info">
<b>Instruction:</b> The `+` operator is called <b>concatenation</b> when used with strings. Explain what is concatenation in your own words.
</div>

YOUR ANSWER HERE

## 2. Basic operators ...

### ... for numbers

Addition

```python
2+3
```

Multiplication

```python
2*3
```

Substraction

```python
2-3
```

Power of a number

```python
3**2
```

Lower than

```python
2 < 3
```

Greater than

```python
2 > 3
```

Equal to

```python
2 == 2
```

Greater/Lower or equal to

```python
print(2 <= 2)
print(2 <= 3)
print(2 >= 3)
```

### ... for strings

Upper case

```python
"Hello world!".upper()
```

Lower case

```python
"Hello world!".lower()
```

Concatenation

```python
"Hello " + "world!"
```

Repetition

```python
"Hello world! " * 3
```

### ... for booleans

Logical Operator | Description
--- | ---
and | If both the operands are True then condition becomes True.
or | If any of the two operands are True then condition becomes True. 
not | Used to reverse the logical (not False becomes True, not True becomes False)

```python
True and False
```

```python
True or False
```

```python
not True
```

<!-- grade_id: cell-4768a22d8ee7e324, points: 0, grade: true, solution: true -->
<div class="alert alert-info">
<b>Instruction:</b> What are the type and values of the following expressions ?
</div>

```python
3 < 2
```


YOUR ANSWER HERE

```python
not False
```

YOUR ANSWER HERE

```python
(3 < 2) or (not False)
```

YOUR ANSWER HERE

## 3. Variables and assignments

In Python, you can assign a value to a variable using the `=` operator. 

A variable is a container that stores a value.

<div class="alert alert-warning">
⚠️ Be careful not to confuse the `=` operator with the `==` operator. The first one is used to assign a value to a variable, while the second one is used to check if two values are equal.
</div>

A variable can be named as you want, but there are some rules to follow:

- Variable names must start with a letter or an underscore `_`,
- Variable names can only contain letters, numbers and underscores,
- Variable names are case sensitive.

<div class="alert alert-info">
<b>Instruction:</b> Create a variable named `my_variable` and assign it the value 42 (int).
</div>

<!-- grade_id: cell-8108190acb4785a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e7567a9ae4b3b7c6, points: 1, grade: true, solution: false -->
```python

```

Now you assigned a value to the variable `my_variable` variable, you can use it in your code.

<div class="alert alert-info">
<b>Instruction:</b> Display the value of `my_variable` and its type.
</div>

<!-- grade_id: cell-d3a28becf6d4f713, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-info">
<b>Instruction:</b> Create a variable named my_variable_2 and assign it the value of my_variable multiplied by 2.
</div>

<!-- grade_id: cell-e38275fd2d0a8170, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-warning">
⚠️ It is common that students get confused with variables in Jupyter Notebook, because code cells can be run in any order, and even multiple times. Look at the example below and explain why the value of `my_variable_3` is changing every time you execute the cell.
</div>

```python
# Initialization
u_n = 0
```

```python

u_n = 5 + u_n
print(u_n)

```

Now you know how to assign, retrieve, and use a value that was stored into a variable.

## 4. Control structures

### 4.1 Conditional structures

In Python, you can use the `if` statement to execute a block of code only if a condition is `True`.

```python
if condition:
    # code to execute if condition is True
```

You can also use the `else` statement to execute a block of code if the condition is `False`.

```python
if condition:
    # code to execute if condition is True
else:
    # code to execute if condition is False
```

Finally, you can use the `elif` statement to execute a block of code if the first condition is `False` and the second condition is `True`.

```python
if condition_1:
    # code to execute if condition_1 is True
elif condition_2:
    # code to execute if condition_2 is True
else:
    # code to execute if condition_1 and condition_2 are False
```

<div class="alert alert-warning">
⚠️ In python, code blocks are defined by their indentation. An indentation is made of 4 spaces, and you can use the `Tab` key to create one.
As a rule of thumb, you should always ident your code after a `:`.
If the identation is not respected, you will get an `IndentationError`.
</div>

<div class="alert alert-info">
<b>Instruction:</b> Create a variable named `my_age` and assign it the value of your age (in years). Then, create a conditional structure that displays "You are eligible to a youth discount" if your age is lower or equal to 26, display "You are not eligible to a discount" if your age is greater than 26 and lower or equal to 65, and display "You are eligible to a senior discount" if your age is greater than 65. 

<!-- grade_id: cell-3f2214963d09110d, points: 1, grade: true, solution: true -->
```python
# Create my_age variable and assign your age to it
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3805e6169aa62780, points: 2, grade: true, solution: true -->
```python
# Create the conditional structure below
# YOUR CODE HERE
raise NotImplementedError()
```

### 4.2 Types of data structure

In Python, you can use different types of data structures to store multiple values.

### 4.2.1 Lists

A list is a collection of values that can be of different types. You can create a list by using square brackets `[]` and separating values with commas `,`.

```python
my_list = ["Hello", "world", 42, 3.14, True]
```

You can access a value in a list by using its index. The index of a list starts at 0.

```python
my_list[0]
```

```python
my_list[4]
```

```python
my_list[5]
```

<!-- grade_id: cell-4de4d365db4b8a24, points: 1, grade: true, solution: true -->
<div class="alert alert-info">
<b>Instruction:</b> Why my_list[5] raised an error ?
</div>

YOUR ANSWER HERE

You can also use negative indexes to access values from the end of the list.

```python
my_list[-1]
```

The `len()` function allows you to know the length of a list.

```python
len(my_list)
```

### 4.2.2 Dictionaries

A dictionary is a collection of key-value pairs. You can create a dictionary by using curly brackets `{}` and separating key-value pairs with commas `,`. The key and the value are separated by a colon `:`.

You can access a value in a dictionary by using its key, that can be numbers or strings: `my_dict[key]`.

In the example below, let's create a fake telephone numbers' dictionary.

```python
# Create a list
my_dict = {}

## Assign a key-value pair to the dictionary
my_dict["John"] = "+49 123456789"
my_dict["Jane"] = "+49 987654321"
```

Display the dictionary using the `print()` function.

<!-- grade_id: cell-6a288a78835435ce, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<div class="alert alert-info"
<b>Instruction:</b> Access the value of the key "Jane" and check if its type is a string using the type() function and the `==` operator.
</div>

<!-- grade_id: cell-4c137e09861f09bd, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now you know two types of data structures and can be used to store and retrieve multiple values.

### 4.3 Loops

In Python, you can use loops to execute a block of code multiple times. There are two types of loops: `for` and `while`.

### 4.3.1 For loop

You can use a `for` loop to iterate over a sequence of values. The syntax is the following:

```python
for value in sequence:
    # code to execute for each value in the sequence
```

In the example below, let's create a list of names and display each name using a `for` loop.


```python
names = ["Alice", "Bob", "Charlie", "David", "Eve"]

for name in names:
    print(name)
```

The function `range()` is very useful to create a sequence of numbers. You can use it with one, two, or three arguments.

```python
range(stop) # create a sequence from 0 to (stop - 1)
range(start, stop) # create a sequence from start to (stop - 1)
range(start, stop, step) # create a sequence from start to (stop - 1) with a step
```

```python
for i in range(5):
    print(i)
```

```python
for i in range(3, 8):
    print(i)
```

<div class="alert alert-info">
<b>Instruction:</b> Create a `for` loop that displays all <b>even numbers</b> comprise between 0 and 100 included, in ascending order.
</div>

<!-- grade_id: cell-0236c1d73a206c2a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-info">
<b>Instruction:</b> Create a `for` loop that displays numbers that are multiples of 3 and comprise between 0 and 100 included, in descending order.
</div>

<!-- grade_id: cell-a3017beccb95f998, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### 4.3.2 While loop

You can use a `while` loop to execute a block of code as long as a condition is `True`. The syntax is the following:

```python
while condition:
    # code to execute as long as the condition is True
```

A `while` loop can be strictly equivalent to a `for` loop, as shown in the example below, but it is less used in Python. 

```python
i = 0
while i < 5:
    print(i)
    i += 1
```

<div class="alert alert-info">
<b>Instruction:</b> Create a `while` loop that displays all <b>even numbers</b> comprise between 0 and 100 included, in ascending order.
</div>

<!-- grade_id: cell-7c340dc3fd86c33f, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-info">
<b>Instruction:</b> Create a `while` loop that displays all numbers that are multiples of 3 and comprise between 0 and 100 included, in descending order.
</div>

<!-- grade_id: cell-43d8b947a5f5fcfb, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-info">
<b>Instruction:</b> Below is a list of names. Create a `while` loop that iterate over the list and display each name until the name "Jane" is found.
</div>

```python
names = ["Alice", "Bob", "Charlie", "David", "Eve", "Frank", "Grace", "Helen", "Isaac", "Jane", "Karl", "Linda", "Mary", "Nathan", "Olivia", "Peter", "Quinn", "Rachel", "Steve", "Tom", "Ursula", "Victor", "Wendy", "Xavier", "Yvonne", "Zach"]
```

```python
for i in range(len(names)):
    print(names[i])
    if names[i] != "Jane":
        print(names[i])
        break
```

<!-- grade_id: cell-d4f9af6a7332cadf, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
for v in names:
    print(v)
    if v == "Jane":
        break
```

```python
i = 0
while names[i] != "Jane":
    print(names[i])
    i += 1
print(names[i])
```

## 5. Functions

In Python, you can define a function to execute a block of code. A function encapsulates a fragment of code into a variable that can be called multiple times.

You can define a function by using the `def` keyword followed by the function name and its parameters, also called arguments. The function name must follow the same rules as variable names.

```python
def my_function(parameter_1, parameter_2, ...):
    # code to execute
    return value
```

Let's see an example of a function that returns the sum of two numbers.


```python
def add(x, y):
    return x + y
```

After defining a function, you can call it by using its name and passing arguments.

```python
add(4, 3)
```

<!-- grade_id: cell-e113d689a177e910, points: 1, grade: true, solution: true -->
<div class="alert alert-info">
<b>Instruction:</b> What is the type of the function's arguments? What is the type returned by the function?
</div>

YOUR ANSWER HERE

Types can be optionally specified in Python, but it is a good practice to do so. You can specify the type of the arguments and the return value by using the `->` operator.

Using the example above, the header would become:

```python
def add(x: int, y: int) -> int:
    # code to execute
    return value
```

It is also a good practice to document your function by using a docstring. A docstring is a string that describes what the function does. It is placed between triple quotes `"""`.

```python
def add(x: int, y: int) -> int:
    """
    This function returns the sum of two integers.

    Parameters:
        x (int): the first integer
        y (int): the second integer
    
    Returns:
        int: the sum of x and y
    """
    # code to execute
    return value
```

<div class="alert alert-info">
<b>Instruction:</b> Create a function named `findNameIndex` using the docstring above.
</div>

```python
"""
This function takes a list of names and a name as arguments and returns the index of the name in the list if it is found, otherwise it returns -1.

Parameters:
    names (list): a list of names
    name (str): the name to find in the list

Returns:
    int: the index of the name in the list if it is found, otherwise -1
"""
```


<!-- grade_id: cell-fd8b0bc3d4285688, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
findNameIndex(names, "Ursula")
```

`assert` is a useful keyword to test functions. It raises an `AssertionError` if the condition is `False`, and nothing if the condition is `True`.

Your fonction `findNameIndex` should not raise any `AssertionError` with the following tests.

```python
names = ["Alice", "Bob", "Charlie", "David", "Eve", "Frank", "Grace", "Helen", "Isaac", "Jane", "Karl", "Linda", "Mary", "Nathan", "Olivia", "Peter", "Quinn", "Rachel", "Steve", "Tom", "Ursula", "Victor", "Wendy", "Xavier", "Yvonne", "Zach"]

assert findNameIndex(names, "Alice") == 0
assert findNameIndex(names, "Bob") == 1
```

<div class="alert alert-info">
<b>Instruction:</b> Write two extra tests that your function should pass.
</div>

<!-- grade_id: cell-0908b9536292f9d6, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 6. Application: british artists and songs

In this section, you will create functions to manipulate a dictionnary of artists and their songs.
The questions will help you to recall the notions seen above.

Let's start by creating a non-exhaustive dictionnary of british artists and their songs.

```python
brit_artists = {"The Beattles" : ["Yesterday",
                            "Hey Jude",
                            "Let it be",
                            "Come Together"],
            "Queen" : ["Bohemian Rhapsody",
                     "We Will Rock You",
                     "We Are The Champions",
                     "Another One Bites The Dust"],
            "Elton John" : ["Rocket man",
                            "Your Song",
                            "Tiny Dancer",
                            "Crocodile Rock"],
            "David Bowie" : ["Space Oddity",
                            "Life on Mars",
                            "Heroes",
                            "Starman"],
            "The Rolling Stones" : ["Paint it Black",
                                    "Sympathy for the Devil",
                                    "Gimme Shelter",
                                    "Angie"]}
                            
```

```python
brit_artists["The Beattles"][0]
```

The variable `brit_artists` is a dictionary (`dict`) whose keys are strings (`str`) and values are lists (`list`) or strings.

<div class="alert alert-info">
<b>Instruction:</b> Implement a function named `getSongs` that takes a dictionary of artists and the name of an artist as arguments and returns the list of songs of the artist if it is found, otherwise it raises a `ValueError`.
</div>

<!-- grade_id: cell-1c252c2311895058, points: 2, grade: true, solution: true -->
```python
def getSongs(data : dict, artist_name : str) -> list:
    """ Get the list of songs from a given artist

    Args:
        data (dict): Dictionary with artists as keys and songs as values
        artist_name (str): Name of the artist

    Raises:
        ValueError: If the artist is not found in the dictionary

    Returns:
        list: List of songs from the artist
    """
    if artist_name in data:
        # YOUR CODE HERE
        raise NotImplementedError()
    else:
        raise ValueError("Artist not found")
```

<div class="alert alert-info">
<b>Instruction:</b> Implement a function named `getArtist` that takes a dictionary of artists and the name of a song as arguments and returns the artist that composed the song if it is found, otherwise it raises a `ValueError`.

Do not forget to indicate the types of the arguments and the return value in the function's header, as well as to write a docstring.
</div>

<!-- grade_id: cell-cb1da99495700477, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b936408aa64684b8, points: 2, grade: true, solution: false -->
```python
assert getSongs(brit_artists, "Queen") == ["Bohemian Rhapsody",
                        "We Will Rock You",
                        "We Are The Champions",
                        "Another One Bites The Dust"]

assert getSongs(brit_artists, "David Bowie") == ["Space Oddity",
                            "Life on Mars",
                            "Heroes",
                            "Starman"]
assert getArtist(brit_artists, "Bohemian Rhapsody") == "Queen"
assert getArtist(brit_artists, "Space Oddity") == "David Bowie"
assert getArtist(brit_artists, "Yesterday") == "The Beattles"
```

## Well done!

You've learned the basics of Python, a powerful programming language that is widely used in data science and machine learning.

In particular, you should now be able to:

 - Identify the main basic data types in Python
 - Use basic operators for numbers, strings, and booleans
 - Assign values to variables using the `=` operator
 - Create conditional structures using `if`, `elif`, and `else`
 - Use lists and dictionaries to store multiple values
 - Use `for` and `while` loops to execute a block of code multiple times or iterate over a sequence of values
 - Define and document functions to encapsulate a fragment of code that can be called later.
 
It is time to learn about data science! Go to the next sheet [Data science tutorial](03_tutorial_data_science.ipynb)